---
## Q2 (6 points)
จงเขียนโปรแกรมที่รับรหัสนิสิต และ ภาคการศึกษา (1/2/3) แล้วแสดงค่าธรรมเนียมการศึกษาของนิสิตในภาค <br>
นั้นตามตารางข้างล่างนี้ซึ่ง เก็บเป็น **dictionary** <br>

![Lab11_2.png](attachment:Lab11_2.png)

เลขหลักที่ 3 (จากหน้า) ของรหัสนิสิตระดับปริญญาบัณฑิต คือ 3, 4 และรหัสนิสิตระดับบัณฑิตศึกษา เป็น 7 <br>
เลข 2 หลักสุดท้ายของรหัสนิสิตเป็นรหัสคณะ <br>
รหัสคณะในกลุ่ม 1 คือ 21 23 25 28 30 31 32 33 35 36 37 38 39 53  <br>
รหัสคณะในกลุ่ม 2 คือ 22 24 26 27 29 34 40 51  <br>
และแสดงผลลัพธ์ดังตัวอย่างข้างล่างนี้ <br>

จงเขียนฟังก์ชัน `get_dict()` ที่สร้าง dictionary จากตารางค่าเทอม ซึ่งมี <br>

- key เป็นสตริงของรหัสคณะในกลุ่ม 1 และ 2 มีค่าเป็น `group1` หรือ `group2`
- value เป็น dictionary ที่มี 
    - key เป็นสตริงของระดับการศึกษา มีค่าเป็น `Bachelor` หรือ `Graduate`
    - value เป็น dictionary ที่มี
        - key เป็นสตริงของภาคการศึกษา ภาค1/2 มีค่าเป็น `S1_2` หรือ ภาค 3 มีค่าเป็น `S3`
        - value เป็น dictionary ที่มี
            - key เป็นสตริงของรหัสนิสิต มีค่าเป็น `ID_1` , `ID_2` หรือ `ID_3` สำหรับรหัสนิสิตที่ขึ้นต้นด้วยปี 48-49 , 50-55 และ 56 เป็นต้นไป ตามลำดับ
            - value เป็นค่าเทอม

และแสดงผลลัพธ์เป็นข้อมูลใน dictionary ตามตัวอย่างการทำงาน


### ตัวอย่างการทำงาน: 

Enter student ID : **6033235423**<br>
Enter semester : **1**<br>
Registration fee : 21000

### ตัวอย่างการทำงาน: 
Enter student ID : **6070013223**<br>
Enter semester : **2**<br>
Registration fee : 31000

### ตัวอย่างการทำงาน: 

Enter student ID : **5473200022**<br>
Enter semester : **3**<br>
Registration fee : 7000

### ตัวอย่างการทำงาน: 

Enter student ID : **5473200022**<br>
Enter semester : **0**<br>
Invalid semester

### ตัวอย่างการทำงาน: 

Enter student ID : **60499**<br>
Invalid ID

### ตัวอย่างการทำงาน: 

Enter student ID : **9913235499**<br> 
Invalid ID

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def get_dict():
    '''สร้าง dictionary ของค่าเทอม แล้วคืนค่า dictionary'''
    return {
        "group1": {
            'Bachelor': {
                'S1_2': {
                    'ID_1': 16000, 
                    'ID_2': 18000, 
                    'ID_3': 21000
                }, 
                'S3': {
                    'ID_1': 4000, 
                    'ID_2': 4500,
                    'ID_3': 5250
                }
            }, 'Graduate': {
                'S1_2': {
                    'ID_1': 22500, 
                    'ID_2': 26000, 
                    'ID_3': 31000
                }, 
                'S3': {
                    'ID_1': 6000, 
                    'ID_2': 7000,
                    'ID_3': 7750
                }
            }
        },
        "group2": {
            'Bachelor': {
                'S1_2': {
                    'ID_1': 12000, 
                    'ID_2': 14500, 
                    'ID_3': 17000
                }, 
                'S3': {
                    'ID_1': 4000, 
                    'ID_2': 4500,
                    'ID_3': 5250
                }
            }, 'Graduate': {
                'S1_2': {
                    'ID_1': 16500, 
                    'ID_2': 19000, 
                    'ID_3': 23000
                }, 
                'S3': {
                    'ID_1': 6000, 
                    'ID_2': 7000,
                    'ID_3': 7750
                }
            }
        }
    }
def run():
    '''รับเลขประจำตัวนิสิต ตรวจสอบ และแสดงค่าเทอม'''
    ID = input('Enter student ID: ')
    if len(ID)!=10 :
        return print ('Invalid ID')
    elif int(ID[:2]) not in range(48,67):
        return print ('Invalid ID')
    elif ID[2] not in "347":
        return print("Invalid ID")
    elif ID[8] != "2" or ID[9] not in "12345678":
        return print("Invalid ID")
    
    semester = input('Enter semester: ')
    if semester not in '123':
        return print('Invalid semester')
    
    semester = 'S3' if semester == '3' else 'S1_2' 
    tier = 'Bachelor' if ID[2] in '34' else 'Graduate'
    group = 'group1' if  ID[8:10] in '21 23 25 28 30 31 32 33 35 36 37 38 39 53'.split() else 'group2'
    year = int(ID[:2])
    if year <= 49:
        yearclass = 'ID_1'
    elif year <= 55:
        yearclass = 'ID_2'
    else:
        yearclass = 'ID_3'
        
    #print (group,tier,semester,yearclass)
    fee = get_dict()[group][tier][semester][yearclass]    
    print (f'Registration fee : {fee}')
```

```python
'''ทดสอบ การสร้าง dict'''
get_dict()
```

Output:
```
{'group1': {'Bachelor': {'S1_2': {'ID_1': 16000, 'ID_2': 18000, 'ID_3': 21000},
   'S3': {'ID_1': 4000, 'ID_2': 4500, 'ID_3': 5250}},
  'Graduate': {'S1_2': {'ID_1': 22500, 'ID_2': 26000, 'ID_3': 31000},
   'S3': {'ID_1': 6000, 'ID_2': 7000, 'ID_3': 7750}}},
 'group2': {'Bachelor': {'S1_2': {'ID_1': 12000, 'ID_2': 14500, 'ID_3': 17000},
   'S3': {'ID_1': 4000, 'ID_2': 4500, 'ID_3': 5250}},
  'Graduate': {'S1_2': {'ID_1': 16500, 'ID_2': 19000, 'ID_3': 23000},
   'S3': {'ID_1': 6000, 'ID_2': 7000, 'ID_3': 7750}}}}
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
assert get_dict()['group1'] == {'Bachelor': {'S1_2': {'ID_1': 16000, 'ID_2': 18000, 'ID_3': 21000}, 'S3': {'ID_1': 4000, 'ID_2': 4500, 'ID_3': 5250}}, 'Graduate': {'S1_2': {'ID_1': 22500, 'ID_2': 26000, 'ID_3': 31000}, 'S3': {'ID_1': 6000, 'ID_2': 7000, 'ID_3': 7750}}}
assert get_dict()['group1']['Bachelor'] == {'S1_2': {'ID_1': 16000, 'ID_2': 18000, 'ID_3': 21000}, 'S3': {'ID_1': 4000, 'ID_2': 4500, 'ID_3': 5250}}
assert get_dict()['group1']['Bachelor']['S3'] == {'ID_1': 4000, 'ID_2': 4500, 'ID_3': 5250}
print('OK')
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
Enter student ID: 5273235421
Enter semester: 2
Registration fee: 18000
```

```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO

def test(inputs,expected_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
def test_easy_print(inputs,expected_output):
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert expected_output in mock_stdout.getvalue(), mock_stdout.getvalue()
```

<!-- grade_id: cell-be181c3483ec7e11, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบการแสดงผลลัพธ์อย่างง่าย'''
test_easy_print(inputs=['5273235421','2'],expected_output='R')
print('OK')
```

<!-- grade_id: cell-25025451c4cc0722, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''ทั่วไป'''
test(['5273235421','2'],'Registration fee : 26000\n')
test(['5273235423','3'],'Registration fee : 7000\n')  
test(['6070013223','1'],'Registration fee : 31000\n') 
test(['6070013228','3'],'Registration fee : 7750\n') 
print('OK')
```

<!-- grade_id: cell-a7595ca4f976893e, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Invalid cases'''
test(['5443200022','0'],'Invalid semester\n')  #
test(['5443200022','4'],'Invalid semester\n')  #
test(['6433235460'],'Invalid ID\n')
test(['6453235430'],'Invalid ID\n') 
test(['4773235430'],'Invalid ID\n')
test(['60499'],'Invalid ID\n') #
print('OK')
```

<!-- grade_id: cell-3478c02530c40be9, points: 1, grade: true, solution: false -->
```python
'''เทสลับ1'''
```

<!-- grade_id: cell-8749df472d842678, points: 1, grade: true, solution: false -->
```python
'''เทสลับ2'''
```